# Lab05: Image Classification

![MNIST data](https://upload.wikimedia.org/wikipedia/commons/thumb/2/27/MnistExamples.png/220px-MnistExamples.png)

```python
%pip install sklearn matplotlib seaborn
```

## TASK

Load the MNIST handwritten digits dataset from sklearn. Build a classifier that can predict the correct number.

### Loading the dataset

```python
from sklearn import datasets
digits = datasets.load_digits()
print(digits.DESCR)
```

### ...  and inspect it
- What object type is digits? Store its type into variable `mnist_type`.
- What are the attributes of the dataset?  How can you access them? Have a look at them and store the individual elements of the digits object stores to the variable `mnist_attributes`.

<!-- grade_id: cell-6d76444aa3818227, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-06eed1b220f62af3, points: 5, grade: true, solution: false -->
```python
# Ignore this cell - this is for automatic testing only.
```

### Use imshow() to plot an image array
Plot one single or several digits. Include their labels in the title and deactivate all axis ticks etc.

<!-- grade_id: cell-b090761cd3341034, points: 7.5, grade: true, solution: true -->
```python
import matplotlib.pyplot as plt
img_nr = 0  # choose an image

# YOUR CODE HERE
raise NotImplementedError()
```

### Preprocessing

With this task, we do not need to much of a preprocessing. But..

A SVC can not work with 2D data. So the data needs to be reshaped to a one-dimensional representation. Store the result to the variable `data`.

<!-- grade_id: cell-21ff99c0826d13b9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8ae02517cb4472d8, points: 7.5, grade: true, solution: false -->
```python
# Ignore this cell - this is for automatic testing only.
```

### Train-Test Split
Split the data into train and test sets (Hint: you can pass data.images AND data.targets simultaneously)

<!-- grade_id: cell-eaccca3abcb3bb6f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-32129f8c633c8f78, points: 10, grade: true, solution: false -->
```python
# Ignore this cell - this is for automatic testing only.
```

### Use a Support Vector Machine as Classifier
Sklearn provides a Support Vecotor Machine (SVM). In particular, next you are using a SVM Classifier (SVC). Import and use the [LinearSVC classifier from sklearn](https://scikit-learn.org/stable/modules/generated/sklearn.svm.LinearSVC.html). 


Do a short research on SVMs.

```python
from sklearn.svm import SVC, LinearSVC
from sklearn.metrics import accuracy_score, confusion_matrix
```

#### Create a SVC classifier instance
Do not train here yet.

<!-- grade_id: cell-a951fd51032ceeff, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c4042a484f4d1a56, points: 10, grade: true, solution: false -->
```python
# Ignore this cell - this is for automatic testing only.
```

#### Fit the classifier on the train set

Check sklearn on how to do that.

<!-- grade_id: cell-c1733f30fd21e17f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-73d8fd21d8cf1bf9, points: 10, grade: true, solution: false -->
```python
# Ignore this cell - this is for automatic testing only.
```

#### Predict the labels for the test set

Store the result to the variable `predicted`. Check sklearn on how to do that.

<!-- grade_id: cell-5a17e39020cf60ad, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-efcc810424b8beb3, points: 10, grade: true, solution: false -->
```python
# Ignore this cell - this is for automatic testing only.
```

### Evaluation
Calculate the accuracy score and store it to the variable `svc_accuracy`.


_Hint:_ have a look on the imports

<!-- grade_id: cell-4e022994d617d3c9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3c0cc4b38c935054, points: 10, grade: true, solution: false -->
```python
assert svc_accuracy > 0.9
```

#### Print the confusion matrix
If you can not recall what a confusion matrix is, have a look at Lecture 5.

<!-- grade_id: cell-e850bd932592b859, points: 10, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Recall and Summerize
So far you did what was asked. With most of the tasks it was explicitly stated what to do. With the following, try to summarize what you did here on this task. Include:
- a description on the task
- a description of the method (SVC: what is it? only overview.)
- state the results and interpret them.

<!-- grade_id: cell-b1785d8ebaf10770, points: 20, grade: true, solution: true -->
YOUR ANSWER HERE